# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 5.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def longest(strings):
    """(list) -> int
    This function finds the longest string in the list of input strings
    and returns the position number of that longest string.

    Parameters
    strings (list: A list of strings.

    Returns
    int: The index of the longest string in the list.

    >>> longest(['happy', 'birthday', 'to', 'me'])
    1

    >>> longest(['one', 'two', 'three', 'four'])
    2
    """

    longest_pos = 0
    longest_len = 0

    print("Position  Item         Longest So Far")
    print("--------- ------------ --------------")
    for position, item in enumerate(strings):
        if len(item) > longest_len:
            longest_pos = position
            longest_len = len(item)
        print("{0:9} {1:12} {2} : {3}".format(position, item, longest_pos, longest_len))

    return longest_pos

### END SOLUTION
```

```python
def longest(my_list):
    """(list)->int
    Find the longest item in a list of strings and return the index for that string.
    If there are multiple strings with the same length, return the index of the first of those longest strings.

    Raises:
    TypeError: If the input is not a list.
    ValueError: If the list is empty or contains no valid strings.

    >>> longest(['happy', 'birthday', 'to', 'me'])
    1

    >>> longest(['one', 'two', 'three', 'four'])
    2
    """

    if not isinstance(my_list, list):
        raise TypeError("That isn't a list!")

    if len(my_list) == 0:
        raise ValueError("There were no items in that list!")

    max_length = -1
    max_index = None

    for num, item in enumerate(my_list):
      try:
         if len(item) > max_length:
             max_length = len(item)
             max_index = num
      except TypeError:
            pass
            if max_index is None:
                raise ValueError("There were no valid strings in that list!")
    return max_index

```

```python
longest(['happy', 'birthday', 2, 'me'])
```

Output:
```
1
```

```python
longest(['happy', 'birthday', 'to', 'me'])
```

Output:
```
1
```

```python
longest(['one', 'two', 123, 'four'])
```

Output:
```
3
```

```python
longest(['abc', 'def', 'ghi'])
```

Output:
```
0
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
PYDEV DEBUGGER WARNING:
sys.settrace() should not be used when the debugger is being used.
This may cause the debugger to stop working correctly.
If this is needed, please check: 
http://pydev.blogspot.com/2007/06/why-cant-pydev-debugger-work-with.html
to see how to restore the debug tracing back correctly.
Call Location:
  File "/usr/lib/python3.10/doctest.py", line 1501, in run
    sys.settrace(save_trace)

Finding tests in NoName
Trying:
    longest(['happy', 'birthday', 'to', 'me'])
Expecting:
    1
ok
Trying:
    longest(['one', 'two', 'three', 'four'])
Expecting:
    2
ok
```

### 5.2 Celsius to Farenheit

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION
def f_to_c(fahrenheit):
    """
    Convert degrees Fahrenheit to degrees Celsius.

    >>> f_to_c(212)
    100.0
    >>> f_to_c(32)
    0.0
    >>> f_to_c(98.6)
    37.0
    """
    return 5/9 * (fahrenheit - 32)

### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
f_to_c(212)
```

Output:
```
100.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

### 5.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION
def los(start_date, end_date, diagnosis):
    """
    Calculate the Length of Stay (LOS) for a patient,
    if the diagnosis is "Observation" then the LOS is 0,
    otherwise it is the difference in days between the start and end dates.

    >>> los(date(2019,1,3), date(2019,1,8), "COPD")
    5
    >>> los(date(2020,1,5), date(2020,1,8), "Observation")
    0
    """
    if diagnosis == "Observation":
        return 0
    return (end_date - start_date).days

encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

for encid, visit in encounters.items():
    print(los(visit['admit'],visit['discharge'],visit['diagnosis']))

# Doctests
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
### END SOLUTION
```

Output:
```
5
4
1
0
0
Finding tests in NoName
Trying:
    los(date(2019,1,3), date(2019,1,8), "COPD")
Expecting:
    5
ok
Trying:
    los(date(2020,1,5), date(2020,1,8), "Observation")
Expecting:
    0
ok
```

```python
assert los(date(2019,1,3), date(2019,1,8), "COPD") == 5
assert los(date(2020,1,5), date(2020,1,8), "Observation") == 0
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    los(date(2019,1,3), date(2019,1,8), "COPD")
Expecting:
    5
ok
Trying:
    los(date(2020,1,5), date(2020,1,8), "Observation")
Expecting:
    0
ok
```

```python
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for encid, visit in encounters.items():
    print(los(visit['admit'],visit['discharge'],visit['diagnosis']))
```

Output:
```
5
4
1
0
0
```

---
---

### 5.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python
test1 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

test2 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

### BEGIN SOLUTION
from datetime import date

def average_los(encounters):
    total_los = 0
    count = 0

    for encounter in encounters.values():
        admit_date = encounter.get("admit")
        discharge_date = encounter.get("discharge")

        if admit_date and discharge_date:
            los = (discharge_date - admit_date).days
            if los > 0:
                total_los += los
                count += 1

    if count == 0:
        return 0  # Return 0 if no valid LOS to average
    return total_los / count



### END SOLUTION
```

```python
from datetime import date

def los(start_date, end_date, diagnosis):
    """
    Calculate the Length of Stay (LOS) for a patient,
    if the diagnosis is "Observation" then the LOS is 0,
    otherwise it is the difference in days between the start and end dates.
    """
    if diagnosis == "Observation":
        return 0
    return (end_date - start_date).days

def average_los(encounters):
    """
    Calculate the average Length of Stay (LOS), not considering individual encounters
    with a LOS of 0.

    >>> round(average_los(encounters),2)
    3.33
    """
    total_days = 0
    total_encounters = 0

    for encid, visit in encounters.items():
        los_value = los(visit['admit'],visit['discharge'],visit['diagnosis'])
        if los_value > 0:
            total_days += los_value
            total_encounters += 1

    return total_days / total_encounters if total_encounters else 0

# Given Encounters data
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

# Perform Assertion Test
assert(round(average_los(encounters), 2) == 3.33)

# Run Doctests
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

```

Output:
```
Finding tests in NoName
Trying:
    round(average_los(encounters),2)
Expecting:
    3.33
ok
```

```python
from datetime import date
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(round(average_los(encounters),2) == 3.33)
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    round(average_los(encounters),2)
Expecting:
    3.33
ok
```

---

### 5.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION
def c_to_f(celsius):
    """
    Convert degrees Celsius to degrees Fahrenheit.

    >>> c_to_f(100)
    212.0
    >>> c_to_f(0)
    32.0
    >>> c_to_f(37)
    98.6
    """
    return (celsius * 9/5) + 32
### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    c_to_f(100)
Expecting:
    212.0
ok
Trying:
    c_to_f(0)
Expecting:
    32.0
ok
Trying:
    c_to_f(37)
Expecting:
    98.6
ok
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.